# <span style="color:teal;">CIS 211 Project 2:  Cryptography</span>

##### Due 11:00 P.M. April 14, 2017

##### Reading:  M&R 3.1 -- 3.3, 3.5; &nbsp; 8.1 -- 8.2, 8.4.1 -- 8.4.4

This week we're going to use simple encryption and decryption algorithms as a way of exploring strings and containers (lists, dictionaries, and sets) in Python.

Encryption is described in Chapter 3 of the textbook.  You should understand how the **substitution cipher** works -- you don't need to understand the details of any of the Python functions that implement the method, but you should know what it means to encode a piece of text with this kind of cipher.

Writing a function that decrypts a piece of text (assuming it really was encrypted with a substitution cipher) is not too hard, compared to other forms of encryption, but still more than we want to do in a one-week project.  Instead of developing a complete algorithm that will decrypt a piece of text we'll write a few small "helper functions" that might be useful as part of a larger application.


###  <span style="color:teal">Imports</span>

```python
from string import punctuation, ascii_uppercase, ascii_lowercase
```

###  <span style="color:teal">Test Data</span>

```python
seuss = '\nUEK XZIV, MOU XZIV,\nRKP XZIV, QHGK XZIV.\n'
print(seuss)
```

Output:
```
UEK XZIV, MOU XZIV,
RKP XZIV, QHGK XZIV.
```

###  <span style="color:teal">Project 1: &nbsp; Ciphertext Words (20 points)</span>

Fill in the body of the function named `cipher_words`.  It will be passed a string containing a "cryptoquote", and it should return the set of words in the quote.  All punctuation should be stripped from the ends of words (but not in the middle).

Example:
```
>>> cipher_words(suess)
{'MOU', 'QHGK', 'RKP', 'UEK', 'XZIV'}
```

**Style (2 points):** This function can be written without a `for` loop, using set comprehension.

##### <span style="color:red">Code</span>

<!-- grade_id: cipher_words, points: 10, grade: true, solution: true -->
```python
def cipher_words(text):
    """
    Extracts words from a given string and converts into a usuable set.

   :param str text: A given set of words as a string.
   :return: The set of words in a str returned as a set.
   :rtype: set
    """
    x = { w.strip(punctuation) for w in text.split() }
    
    return x
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: cipher_words_test, points: 5, grade: true, solution: false -->
```python
# pass the example text to the function, make sure 5 words are returned
words = cipher_words(seuss)
assert len(words) == 5

# look up each word in the result set
for w in ['MOU', 'QHGK', 'RKP', 'UEK', 'XZIV']:
    assert w in words
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: cipher_words_doc, points: 5, grade: true, solution: true -->
Used set comprehension as opposed to a for loop to iteratively strip each part of the created set.

###  <span style="color:teal">Project 2: &nbsp; Count Ciphertext Letters (20 points)</span>

Fill in the body of the `count_unsolved` function.  It will be passed a single word, and it should return the number of unique ciphertext (upper case) letters in the word.

Examples:
```
>>> count_unsolved('XZIV')
4

>>> count_unsolved('QTETET')
3
```
Note how there are 6 letters in the second example, but some are repeated.  There are only 3 different letters:  Q, E, and T.

##### <span style="color:red">Code</span>

<!-- grade_id: count_unsolved, points: 10, grade: true, solution: true -->
```python
def count_unsolved(word):
    """
    Returns the unique ciphertext letters in a word.

   :param str word: A given string that is a single word.
   :return: the number of unique ciphertext (upper case) letters in the word.
   :rtype: integer
    """
    x = len(set("".join(ch for ch in word if ch.isupper())))
    
    return x
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: count_unsolved_test_1, points: 1, grade: true, solution: false -->
```python
# each letter in this string is unique
assert count_unsolved('XZIV') == 4
```

<!-- grade_id: count_unsolved_test_2, points: 1, grade: true, solution: false -->
```python
# there are only 3 unique letters in this string
assert count_unsolved('QTETET') == 3
```

<!-- grade_id: count_unsolved_test_3, points: 1, grade: true, solution: false -->
```python
# repeat the previous test after "solving" one of the letters
assert count_unsolved('QaEaEa') == 2
```

<!-- grade_id: count_unsolved_test_4, points: 1, grade: true, solution: false -->
```python
# a word can be completely solved (no remaining ciphertext letters)
assert count_unsolved('banana') == 0
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: count_unsolved_doc, points: 6, grade: true, solution: true -->
In creating a code for this function I used the .isupper function which checks the uppercase letters in a string and
used the .join function to create a new string of only capital letters. I converted this string to a set to remove any duplicate letters and checked the length of the set.

###  <span style="color:teal">Project 3: &nbsp; Fewest Unsolved Letters (15 points)</span>

One strategy for choosing words to work on is to find the word that has the fewest letters remaining to be solved.  Fill in the body of the `fewest_unsolved` function so it iterates over a set of words and returns a tuple `(n,w)` where `w` is the word with the fewest remaining undecrypted letters and `n` is the number of undecrypted letters.

For example, suppose the quote from earlier has been partially decrypted and the set of words is now
```
{'MOo', 'QHue', 'ReP', 'XiIV', 'oEe'}
```
Passing this set to `fewest_unsolved` should return `oEe` because it has one uppercase letter and the others all have two:
```
```

**Note:** If there is a tie (more than one word has the same minimum number of undecrypted letters) just return one at random (_e.g._ the first one found).

##### <span style="color:red">Code</span>

<!-- grade_id: fewest_unsolved, points: 6, grade: true, solution: true -->
```python
def fewest_unsolved(group):
    """
    Iterates over a set of words and returns the word with the fewest remaining undecrypted letters.

   :param set group: A given set containing multiple strings.
   :return: the string with the fewest remaining undecrypted letters.
   :rtype: string
    """
    unsolved_list = []
    for w in group:
        unsolved_list.append([count_unsolved(w),w])
    x = sorted(unsolved_list)
    try:
        return x[0][1]
    except:
        print ("Set was not given")
    
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: fewest_unsolved_test1, points: 2, grade: true, solution: false -->
```python
# this set has a word with one ciphertext letter and three others with two ciphertext letters
test_set = {'oEe', 'XiIV', 'MOo', 'QHue', 'ReP'}

# the first call should find the word with one ciphertext letter
w = fewest_unsolved(test_set)
assert w == 'oEe'
```

<!-- grade_id: fewest_unsolved_test_2, points: 2, grade: true, solution: false -->
```python
# remove that word from the set, the next call should find a word with two ciphertext letters
test_set.remove(w)
w = fewest_unsolved(test_set)
assert count_unsolved(w) == 2
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: fewest_unsolved_doc, points: 5, grade: true, solution: true -->
The process in creating this function involved:

1.creating an empty list to store the tuples of data.

2.Creating a for loop which called count_unsolved function and added its output and the entered string into a tuple.

3.sorted the list with the fewest remaining undecrypted letters first.

4.used the try and except method in the event that the entered string is an empty set which would cause the program to crash, so instead of crashing, returns the user an error statement.

###  <span style="color:teal">Project 4: &nbsp; Letter Pairs (30 points)</span>

An effective strategy for decrypting English text is to consider letters two at a time (see Section 8.4.3 of the textbook for an explanation of why this approach is helpful).

For this part of the project, fill in the definition of the function named `make_letter_pairs`.  It will be passed a string, and it should return a matrix $m$, such that $m_{i,j}$ is the frequency with which letters $i$ and $j$ were next to each other in the string.

One way to create a matrix in Python is to use a list of lists, but for this project we'll make a dictionary of dictionaries.  If we save the result of the function call in a variable named `m`, we can look up a frequency using indexing.  For example, we can look up the frequency of `'i'` followed by `'e'` by evaluating
```
>>> m['i']['e']
```

Here is an example:
```
>>> m = make_letter_pairs('how now, brown cow?')
>>> m
{'b': {'r': 0.1},
 'c': {'o': 0.1},
 'h': {'o': 0.1},
 'n': {'o': 0.1},
 'o': {'w': 0.4},
 'r': {'o': 0.1},
 'w': {'n': 0.1}}
```

The 10 places in this text where there are two letters next to each other are `ho`, `ow`, `no`, `ow`, `br`, `ro`, `ow`, `wn`, `co`, and `ow`.  The combination of `b` followed by `r` occurs only once, so its frequency is 1 / 10 = 0.1:
```
>>> m['b']['r']
0.1
```

The pair `ow` occurs four times, so its frequency is 4 / 10 = 0.4:
```
>>> m['o']['w']
0.4
```

**Note**  There is a similar function, named `neighborCount`, defined in the book. You can use this function to get some ideas for how to write your code, but there are some important differences you should be aware of:
* the function in the book creates counts, but our function creates frequencies
* the function in the book counts in both directions, _e.g._ when it sees the pair `ow` it adds 1 to the count of `ow` and one to the count for `wo`; we count one direction only

##### <span style="color:red">Code</span>

<!-- grade_id: letter_pairs, points: 15, grade: true, solution: true -->
```python
def make_letter_pairs(text):
    """
    When passed a string, and it should returns a matrix showing the 
    frequency of letters next to each other in the string.

   :param string text: A given string of text.
   :return: a matrix containing the frequency of letters appearing next to each other.
   :rtype: dictionary
    """
    ch_dict = {}
    n=0
    newtext = ''.join(ch for ch in text if ch not in set(punctuation))
    newtext.lower()
    
    for i in range(len(newtext)-1):      #loop that creates matrix of counts
        nblist = ch_dict.setdefault(newtext[i],{})
        if newtext[i+1] in ascii_lowercase:
            nblist[newtext[i+1]] = (nblist.setdefault(newtext[i+1],0)+1)
            ch_dict.pop(" ", None)
            
    for x in ch_dict:          #loop that iterates through dictionary and creates total count of pair frequencies
        for z in ch_dict[x]:
            n += (ch_dict[x][z])
            
    for x in ch_dict:     
        for z in ch_dict[x]:
            ch_dict[x][z]=(ch_dict[x][z]/n)
            
    return ch_dict
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: letter_pairs_test_1, points: 4, grade: true, solution: false -->
```python
# the first test uses the 'brown cow' example
m = make_letter_pairs('how now, brown cow?')

# there are 7 different letters in the string
assert len(m) == 7

# check two of the frequencies
assert round(m['o']['w'],1) == 0.4
assert round(m['r']['o'],1) == 0.1
```

<!-- grade_id: letter_pairs_test_2, points: 2, grade: true, solution: false -->
```python
# the remaining tests are based on frequencies using a large body of text (War of the Worlds, by H. G. Wells)
with open('wells.txt') as f:
    pp = make_letter_pairs(f.read())

# look up some common letter combinations ("pp" stands for "pair probability")
assert round(pp['e']['e'], 3) == 0.006
assert round(pp['s']['s'], 3) == 0.004
```

<!-- grade_id: letter_pairs_test_3, points: 2, grade: true, solution: false -->
```python
# look up some rare combinations (in this book, at least)
assert round(pp['r']['z']) < 0.000001
assert round(pp['y']['n']) < 0.000001
```

<!-- grade_id: letter_pairs_test_4, points: 2, grade: true, solution: false -->
```python
# letter combinations 'jx' and 'kk' never occur in this book
assert 'x' not in pp['j']
assert 'k' not in pp['k']
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: letter_pairs_doc, points: 5, grade: true, solution: true -->

The process in creating this function involved:

1.creating an empty dictionary and removing punctuation and uppercases from the string.

2.Creating a for loop that creates matrix of counts and .pop("") the ""value out of the dict.

3.creating a loop that iterates through dictionary and creates total count of pair frequencies.

4.creating a final loop that iterates though all dict value and divides each entry by total count of pair frequencies.


###  <span style="color:teal">Project 5: `apply` (15 points)</span>

The final part of the project is to write a function named `apply` that will use a decryption key (a dictionary mapping ciphertext letters to plaintext) to update a partially decrypted string.

For example, this key replaces X with `'f'`, Z with `'i'`, I with `'s'`, and V with `'h'`:
```
>>> k = {'X':'f', 'Z':'i', 'I':'s', 'V':'h'}
```

If we apply the key to the example ciphertext this is the result:
```
>>> apply(k, 'UEK XZIV, MOU XZIV, RKP XZIV, QHGK XZIV.')
'UEK fish, MOU fish, RKP fish, QHGK fish.'
```

##### <span style="color:red">Code</span>

<!-- grade_id: apply, points: 6, grade: true, solution: true -->
```python
def apply(key, quote):
    """
    Iterates through each character in a string and replaces character with its decrypted key value.

   :param dict key: A given dict containing a decryption key.
   :param string quote: a encoded string of text.
   :return: a partially decrypted string.
   :rtype: string
    """
    for ch in quote:
        if ch in key:
            quote = quote.replace(ch, key[ch])
            
    return quote
```

##### <span style="color:red">Autograder Tests</span>

<!-- grade_id: apply_test, points: 4, grade: true, solution: false -->
```python
# make a decryption key to use in the tests
k = {'X':'f', 'Z':'i', 'I':'s', 'V':'h'}

# apply the key 
assert apply(k,'XZIV') == 'fish'
assert apply(k,'UDKLCXZIV') == 'UDKLCfish'
assert apply(k,'XMIVZQP') == 'fMshiQP'
assert apply(k,'JSBWC') == 'JSBWC'
```

##### <span style="color:red">Documentation</span>

<!-- grade_id: apply_doc, points: 5, grade: true, solution: true -->
The process in creating this function involved:

1.iterating over each character in the quote using a for loop.

2.checking if the character was a key in the dictionary

3.if it was, the character in the string would be replaced by its matching key value.