# <span style="color:teal;">CIS 211 Project 6:  N-Body GUI</span>

##### Due 11:00 P.M.  May 21, 2017

##### Reading:  Online Tutorials

For this week's project you will create the elements used in a GUI for visualizing the solar system simulation: 
* a subclass of Body that adds attributes used to display a planet 
* a new type of Canvas with methods for drawing and moving planets
* a subclass of Tk's spinbox used to control the number of planets to display
* a frame that has a run button and a place to enter the simulation parameters

###  <span style="color:teal">N-Body Simulation Code</span>

This week's project will use the Vector and Body classes from Project 4 and the Planet class and `step_system` function from Project 5.  There are two ways to use that code in this project:  (1) copy your definitions from the previous projects, or (2) download the instructor's solution from Canvas.

Instructions for both methods are in the next two sections.  **Use either of these methods, but not both.**  

Note: you will earn **extra credit points** if you use your own code.

####  <span style="color:teal">Option 1: &nbsp; Use Your Own Code</span>

Choose this method only if your class definitions and `step_system` functions from the Planets and Orbits projects passed all their unit tests. Copy the complete definitions from your other notebooks and paste them into the code cells below.  

```python
from math import sqrt

class Vector:
    """
    A Vector object represents a 3 dimensional geometric object that has length and direction in space.
    It is represented as the coordninates (x,y,z) which correspond to its location.
    """
    
    def __init__(self,x,y,z):
        self._x = x
        self._y = y
        self._z = z
        
    def __repr__(self):
        return "({:.3g},{:.3g},{:.3g})".format(self._x,self._y,self._z)
    
    def x(self):
        return self._x
    
    def y(self):
        return self._y
    
    def z(self):
        return self._z
    
    def norm(self):
        """
        Computes the magnitude of a vector.

       :param Vector self: A given Vector of coordinates.
       :return: the magnitude of a given vector.
       :rtype: float
        """
        return (sqrt((self._x**2)+(self._y**2)+(self._z**2)))
    
    def __add__(self, other):
        return Vector(self._x + other._x, self._y + other._y, self._z + other._z)
    
    def __sub__(self, other):
        return Vector(self._x - other._x, self._y - other._y, self._z - other._z)
    
    def __mul__(self,other):
        return Vector(self._x * other, self._y * other, self._z * other)

    def __eq__(self,other):
        return Vector((round(self._x-other._x,3) == 0), (round(self._y-other._y,3) == 0) ,round(self._z-other._z,3) == 0)
    
    def __ne__(self,other):
        return Vector((round(self._x-other._x,3) != 0), (round(self._y-other._y,3) != 0) ,round(self._z-other._z,3) != 0)
    
    def clear(self):
        """
        Clears the value of a Vector to all 0's.

       :param Vector self: A given Vector of coordinates.
       :return: a Vector of values with coordinates (0,0,0).
       :rtype: Vector
        
        """
        self._x = 0
        self._y = 0
        self._z = 0
```

```python
G = 6.67E-11

class Body:
    """
    A Body object represents the state of a celestial body.  A body has mass 
    (a scalar), position (a vector), and velocity (a vector).  A third vector, 
    named force, is used when calculating forces acting on a body.
    """
    def __init__(self, mass=0,position=Vector(0,0,0), velocity=Vector(0,0,0), force=Vector(0,0,0)):
        self._mass = mass
        self._position = position
        self._velocity = velocity
        self._force = force
        
    def __repr__(self):
        return "{:.3g}kg {} {}".format(self._mass,self._position,self._velocity)
    
    def mass(self):
        return self._mass
    
    def position(self):
        return self._position
    
    def velocity(self): 
        return self._velocity
    
    def force(self):
        return self._force
    
    def direction(self, other):
        """
        a method which takes another Body object 
        as a parameter and returns a vector that "points at" the other Body.

       :param Body self: A Body object.
       :param Body other: Another Body object.
       :return: a body object with an updated position vector that "points at" the other Body.
       :rtype: Body
        """
        return self._position-other._position
    
    def clear_force(self):
        """
        a method that sets the force vector to (0,0,0) by calling that vector's clear method

       :param Body self: A Body object.
       :return:A body object with the the force updated with the  force vector
       of a Body object to  being reset to (0,0,0).
       :rtype: Body
        
        """
        return self._force.clear()
            
    def add_force(self,other):
        """
        a method that computes a force of an object pulling on the force of another object.

       :param Body self: A Body object.
       :param Body other: Another Body object.
       :return: A body object with the the force updated with the force its is being pulled on by another object.
       :rtype: Body
        
        """
        self._force = self._force - ((self.direction(other)) * ((other._mass) / (self.direction(other).norm()**3)))
        return self._force
    
    def move(self,dt):
        """
        a method that uses the current value of the force vector to update the object's position
        with the objects acceleration

       :param Body self: A Body object.
       :param number dt: a time step size 
       :return: Body object with an updated position vector .
       :rtype: Body
        
        """
        A= self._force * G
        self._velocity = self._velocity+A*dt
        self._position = self._position + self._velocity*dt
```

```python
class Planet(Body):
    """
    A Planet object represents the state of a celestial body.  A body has mass 
    (a scalar), position (a vector), and velocity (a vector).  A third vector, 
    named force, is used when calculating forces acting on a body.
    """
    def __init__(self,mass=0,position=Vector(0,0,0), velocity=Vector(0,0,0), force=Vector(0,0,0), color="", name = ""):
        """
        Create a new Planet with mass, position,velocity,force,color, and name
        
        :param Planet self: A Planet object.
        :param float mass: the mass of a Planet
        :param Vector position: the position of the planet
        :param Vector velocity: the velocity of a planet
        :param Vector force: the force of a planet
        :param string color:the color of a planet
        :param string name: the name of a planet
        :return: Planet object.
        :rtype: Planet
        """
        super().__init__(mass,position, velocity, force)
        self._color= color
        self._name= name
        
    def name(self):
        """
        returns the name of the Planet
        
       :param Body self: A Planet object.
       :return: string explaining the assigned name of the Planet object.
       :rtype: string
        """
        return self._name
    
    def color(self):
        """
        returns the color of the planet
        
       :param Body self: A Planet object.
       :return: string explaining the assigned color of the Planet object.
       :rtype: string
        """
        return self._color
```

```python
def step_system(bodies, dt=86459, nsteps=1):
    """
    Draws a set of concentric circles.  

    :param body bodies: a list of body and planet items
    :param number dt:size of the time step
    :param number nsteps: number of times to simulate
    :return: returns the x,y coordinates of the location at each time step for each planet
    :rtype: List
        
    """
    n = len(bodies)
    orbits = [[] for p in range(n)]
    counter=0
    
    for l in range (nsteps):
        for i in range(n-1):
            for j in range (n-1):
                if i != j:
                    bodies[i].add_force(bodies[j])
            bodies[i].move(dt)
            bodies[i].clear_force()
            
        for body in bodies:
            x=body.position().x()
            y=body.position().y()
            z=(x,y)
            orbits[counter].append(z)
            counter +=1
        counter=0
                
    return orbits
```

####  <span style="color:teal">Option 2: &nbsp; Download `Orbits.pyc` from Canvas </span>

Download one of the `Orbits.cpython-3x.pyc` files, rename it `Orbits.pyc`, and move it to the same folder as this notebook.  Then uncomment and execute the `import` command in the code cell below.

**Make sure you get the right file for your version of Python.**

```python
#from Orbits import *
```

###  <span style="color:teal">Libraries and Constants</span>

Execute these code cells each time you open the notebook.  The cell that defines two global constants will allow us to run autograder tests when the notebook is loaded into Jupyter but not when the program is run from the command line.

```python
IPython = (__doc__ is not None) and ('IPython' in __doc__)
Main    = __name__ == '__main__'
```

```python
if IPython:
    %gui tk
```

```python
import tkinter as tk
import tkinter.ttk as ttk
from tkinter import filedialog
from time import sleep
```

###  <span style="color:teal">Helper Function</span>

To help test and debug your code execute this code cell to define a function named `read_bodies`.

The first argument passed to the function should be the name of the file containing planet descriptions.  Download `solarsystem.txt` from Canvas and save it in the same directory as this notebook.

The second argument is the type of Body object you want to make.  If you call
```
>>> read_bodies('solarsystem.txt', Body)
```
you will get back a list of Body objects (you should to this to make sure you can read the data file).  

Part 1 of this week's project is to defined a new class named TkPlanet that is derived from the Planet class. If you call
```
>>> read_bodies('solarsystem.txt', TkPlanet)
```
you will get a list of TkPlanet objects that include the attributes needed to draw the planets on a canvas.

```python
def read_bodies(filename, cls):
    '''
    Read descriptions of planets, return a list of body objects.  The first
    argument is the name of a file with one planet description per line, the
    second argument is the type of object to make for each planet.
    '''
    if not issubclass(cls, Body):
        raise TypeError('cls must be Body or a subclass of Body')

    bodies = [ ]

    with open(filename) as bodyfile:
        for line in bodyfile:
            line = line.strip()
            if len(line) == 0 or line[0] == '#':
                continue
            name, m, rx, ry, rz, vx, vy, vz, rad, color = line.split()
            args = {
                'mass' : float(m),
                'position' : Vector(float(rx), float(ry), float(rz)),
                'velocity' : Vector(float(vx), float(vy), float(vz)),
            }
            opts = {'name': name, 'color': color, 'size': int(rad)}
            for x in opts:
                if getattr(cls, x, None):
                    args[x] = opts[x]
            bodies.append(cls(**args))

    return bodies

```

###  <span style="color:teal">Part 1: &nbsp; TkPlanet (20 points)</span>

Fill in the definition of the TkPlanet class in the code cell below.  The constructor should accept the same arguments as the Planet class constructor plus one new attribute:
* `size` is the radius, in pixels, of the circle to draw to represent the body

Define a getter method named `size` that will return the value of size attribute.

A second new attribute of a TkPlanet object will hold the ID of the circle created when the planet is drawn. Initialize this attribute to None when the planet is created, and define a "setter" named `set_graphic` and a "getter" named `graphic`.

**Note:** If you need a reminder of the arguments passed to Body type `?Body` or `help(Body)` in a Jupyter code cell.


##### <span style="color:red">Code:</span>

<!-- grade_id: tkbody, points: 10, grade: true, solution: true -->
```python
class TkPlanet(Planet):
    """
    A TKPlanet object represents a updated type of Planet object with the inclusion of a size variable.
    """
    
    def __init__(self,mass=0,position=Vector(0,0,0), velocity=Vector(0,0,0), force=Vector(0,0,0), color="", name = "",size=0):
        """
        Creates a new TKPlanet Object with mass, position,velocity,force,color,and size
        
        :param Planet self: A TKPLanet object.
        :param float mass: the mass of a TKPLanet
        :param Vector position: the position of the TKPLanet
        :param Vector velocity: the velocity of a TKPLanet
        :param Vector force: the force of a TKPLanet
        :param string color:the color of a TKPLanet
        :param string name: the name of a TKPLanet
        :param int size: size in radius of the TKPLanet
        :return: TKPlanet object.
        :rtype: TKPlanet
        """
        super().__init__(mass,position, velocity, force,color,name)
        self._size = size
        self._graphic = None
        
    def set_graphic(self,other):
        """
        a setter function that sets the object ID of a TKplanet object
        """
        self._graphic = other
    
    def graphic(self):
        """
        a getter function that gets the TKplanet object
        """
        return self._graphic
        
    def size(self):
        """
        returns the size in radius of the TKPLanet
        """
        return self._size


```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests of the TkBody class.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

<!-- grade_id: tkbody_test1, points: 2, grade: true, solution: false -->
```python
if IPython:
    b = TkPlanet(0, Vector(0,0,0), Vector(0,0,0), '', '', 0)

    assert isinstance(b, Planet)
    assert isinstance(b, TkPlanet)

    assert b.size() == 0
    assert b.graphic() is None

    b.set_graphic(0)
    assert b.graphic() == 0
```

<!-- grade_id: tkbody_test2, points: 3, grade: true, solution: false -->
```python
if IPython:
    bodies = read_bodies('solarsystem.txt', TkPlanet)

    assert isinstance(bodies[0], TkPlanet)

    assert bodies[0].name() == 'sun'
    assert bodies[0].size() == 10
    assert bodies[0].color() == '#ffff00'
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: tkbody_doc, points: 5, grade: true, solution: true -->
The process in creating this class involved:

1.defining the TKplanet class as a child class inherting from the Planet class.

2.defining getter and setter functions for the graphic attribute of the TKPlanet Class.

3.adding the size attribute for the Class.

###  <span style="color:teal">Part 2: &nbsp; SolarSystemCanvas (25 points)</span>

A SolarSystemCanvas is a type of Canvas used to draw planets.  An instance of this class has the following methods:
* `set_planets` is passed a list of all the bodies in the simulation; it should save this list in an instance variable for use later
* `view_planets` is passed an integer `n`, and it should draw circles representing the first `n` bodies in the list (and save the value of `n` in an instance variable)
* `move_planets` is also passed a complete list of body objects; it should use the positions of the first `n` bodies to compute new locations of the circles of the `n` planets currently on the canvas

Much of the code for this class has been written already for you.  Here are the details of the methods you need to fill in.

#### `_compute_scale` 

An important role for this class is to convert solar system coordinates (billions of meters) into screen coordinates (hundreds of pixels).  The method named `_compute_scale` calculates a scaling factor so that when a body is displayed it is placed at the correct location on the screen.  The formula is

$$f = \frac{p\,/\,2}{| \, d_\mathrm{max} \, |} $$

where $p$ is the height or width (whichever is smaller) of the canvas, in pixels, and $d_\mathrm{max}$ is the largest distance to any planet in the list passed to the method.

For example, suppose the screen is 500 pixels wide and 400 pixels tall, and this method is called with a list of 6 bodies, for the sun and all the planets up through Jupiter.  The largest distance is the $x$ coordinate of Jupiter at the start of the simulation, which is $-7.5 \times 10^{11}$.  Plug these values into the equation to get 

$$
f = \frac{200} { 7.5 * 10^{11} } = 2.67 \times 10^{-9}
$$

There are different ways to figure out the largest distance.  A simple method is to find the maximum of the norm of each position vector; another is to find the largest absolute value of any $x$ or $y$ coordinate.  

You can also include a "fudge factor" if you want to add a little extra space so the furthest orbit does not bump up against the edge of the canvas.  For example, if you use $1.1 \times d_\mathrm{max}$ in the denominator the planets will be about 10% closer to the center.

Your method should compute the scale factor and save it in an instance variable named `_scale` so it can be used later when the planets are drawn and moved.

**Important:** You need to save the scale factor in `self._scale` because that is the name used by the other methods which we have written for you.

#### `view_planets` 

This method will be passed an integer `n`, and it needs to draw circles representing the first `n` planets in the current list of TkPlanet objects (which have been saved in the instance variable `self._planets`).  Your method needs to:
* compute the scale factor for this list of planets
* delete any planets and lines that are on the canvas
* draw circles for the first `n` planets in `self._planets`
* save the value of `n` in `self._outer` so it can be used by `move_planets`

The circles should use the size and color attributes of the TkPlanet objects that represent the planets.

#### `move_planets` 

This method will be called by the main simulator after each time step.  The argument passed to the method is the list of TkPlanet objects with their updated position vectors.  You want to move the circles for the first $n$ of these objects, where $n$ is the number of planets currently displayed on the canvas.

For each planet currently on the canvas:
* get the $x$ and $y$ coordinates of the current location of the circle
* compute the $x$ and $y$ coordinates for the updated location
* move the circle to the new location
* draw a line segment from the old location to the new location

**Note:** Two methods that compute screen locations have been written for you.  
* Call `self._current_loc(p)` to get the $x$ and $y$ coordinates of the center of the circle representing planet `p`.  
* Call `self._compute_loc(p)` to get the $x$ and $y$ coordinates for the new center of the circle based on the position vector of planet `p`.

<!-- grade_id: solar_system_canvas, points: 14, grade: true, solution: true -->
```python
class SolarSystemCanvas(tk.Canvas):
    def __init__(self, parent, height=600, width=600):
        tk.Canvas.__init__(self, parent, height=height, width=width, background='gray90', highlightthickness=0)
        self._planets = None
        self._outer = None
        self._scale = None
        self._offset = Vector(int(self['width'])/2, int(self['height'])/2, 0)
        
    def set_planets(self, lst):
        self._planets = lst
        self._outer = len(lst)
        self._compute_scale(lst)
        self.view_planets(len(lst))
        
    def view_planets(self, n):
        """
        method that when passed an integer n draws circles representing the first n planets 
        in the current list of TkPlanet objects 
        
        :param self
        :param int n:the desired number of planets to be displayed
        :return: None
        :rtype: None
        """
        self._compute_scale(self._planets[0:n])
        for z in self.find_all():
            self.delete(z)
        for planet in self._planets[0:n]:
            (x,y)=self._compute_loc(planet)
            ID = self.create_oval((x-planet.size()),(y-planet.size()),(x+planet.size()),(y+planet.size()),fill=planet.color())
            planet.set_graphic(ID)
        self._outer = n
        
        
    def move_planets(self, lst):
        """
        method that when passed a list of objects, moves the objects on the canvas to their updated locatin, along
        with a line showing their movement
        :param self
        :param list lst:the list of objects that are to be moved and displayed
        :return: None
        :rtype: None
        """
        for planet in lst[0:self._outer]:
            c=planet.graphic()
            to=self._compute_loc(planet)
            from_=self._current_loc(planet)
            self.create_line(from_,to)
            self.move(c,to[0]-from_[0],to[1]-from_[1])
            self.update_idletasks()  
            
    def _compute_scale(self, bodies):
        """
        method that when passed a list of objects, computes the scale factor for which the coordonates for each object 
        can be converted into a viewable canvas values.
        :param self
        :param list TKplanets:a list of given planets
        :return: None
        :rtype: None
        
        """
        p = min((int(canvas.cget("width"))),(int(canvas.cget("height"))))/2
        x = [abs(body.position().x()) for body in bodies]
        y = [abs(body.position().y()) for body in bodies]
        d_max= (max(max(x),max(y)))*1.1
        self._scale = p/d_max
    
    def _compute_loc(self, p):
        pos = p.position() * self._scale + self._offset
        return pos.x(), pos.y()
    
    def _current_loc(self, p):
        ul, ur, _, _ = self.coords(p.graphic())
        return ul + p.size(), ur + p.size()
```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests of the SolarSystemCanvas class.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

```python
if IPython:
    canvas = SolarSystemCanvas(tk.Tk(), height=400, width=400)
    canvas.pack()
    bodies = read_bodies('solarsystem.txt', TkPlanet)
```

<!-- grade_id: canvas_test1, points: 3, grade: true, solution: false -->
```python
if IPython:
    canvas.set_planets(bodies)
    canvas.view_planets(2)
    assert [canvas.type(x) for x in canvas.find_all()].count('oval') == 2
```

<!-- grade_id: canvas_test2, points: 3, grade: true, solution: false -->
```python
if IPython:
    for i in range(10):
        step_system(bodies)
        canvas.move_planets(bodies)
    tk_objects = canvas.find_all()
    assert len(tk_objects) == 22
    assert [canvas.type(x) for x in tk_objects].count('line') == 20
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: canvas_doc, points: 5, grade: true, solution: true -->
The process in creating this class involved:

1.adding to the class by creating the compute scale function which i did by iterating over the x and y position coordinates a finding the largest absolute value and using that with the given algorithm.

2.defining the view planets function which graphs the starting location of the planets using the create oval function and manipulating it into making a circle instead of an oval and saving the ID of each created circle.

3.defining the move planets function which graphs the planets movements from their starting positions by using the saved ID of each planet and making a call to move by using that ID.

###  <span style="color:teal">Part 3: &nbsp; Viewbox (15 points)</span>

A Viewbox is a special type of Spinbox used to specify the number of planets to display on the canvas.  A Viewbox will have a minimum value of 2 (the canvas should always show at least the Sun and Mercury) and a maximum value equal to the number of bodies in the data set.

The `__init__` function has been written for you.  It initializes the new Viewbox by passing the callback function to the parent class constructor (this function will be called when users click either the up or down button in the Spinbox).

Note that a new Viewbox widget is initially disabled.  That's because we won't know how many bodies are in the data set until later, when the data is read from a file.  When the data is loaded, the top level application will call a method named `set_limit`.  You need to fill in the body of this method:
* enable the spinbox (this must be done first, before the other operations)
* set the lower bound to 2, and the upper bound to the number of bodies
* change the value displayed in the spinbox to the number of bodies

##### <span style="color:red">Code:</span>

<!-- grade_id: viewbox, points: 5, grade: true, solution: true -->
```python
class Viewbox(tk.Spinbox):
    def __init__(self, parent, callback):
        tk.Spinbox.__init__(self, parent, command=callback, width=3, state=tk.DISABLED)
                        
    def set_limit(self, nbodies):
        """
        A method that sets the limit for the viewbox from 2 to the value of the given integer nbodies, as well as
        enables the viewbox button.
        """
        self["state"]=tk.NORMAL
        self["to"]= nbodies
        self["from"]= 2
        self.delete(0, tk.END)
        self.insert(0,nbodies)
    
```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests of the Viewbox class.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

```python
def vb_cb():
    global called
    called = True

if IPython:
    vb = Viewbox(tk.Tk(), vb_cb)
    vb.pack()        
```

<!-- grade_id: vb_test_new, points: 1, grade: true, solution: false -->
```python
if IPython:
    called = False
    vb.invoke('buttondown')
    vb.invoke('buttonup')
    assert not called
```

<!-- grade_id: vb_test_set, points: 2, grade: true, solution: false -->
```python
if IPython:
    vb.set_limit(5)
    assert vb.get() == '5'
    assert vb['from'] == 2
    assert vb['to'] == 5
```

<!-- grade_id: vb_test_use, points: 2, grade: true, solution: false -->
```python
if IPython:
    called = False
    vb.invoke('buttondown')
    assert vb.get() == '4'
    vb.invoke('buttonup')
    assert vb.get() == '5'
    assert called
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: view_doc, points: 5, grade: true, solution: true -->
The process in creating this class involved:

1: filling in the set limit function which involved first enabling the spinboxes button, then setting the upper and lower bounds for which it uses to adjust.

###  <span style="color:teal">Part 4: &nbsp; RunFrame (20 points)</span>

A RunFrame is a container that has a run button, text entry boxes for the time step size and number of steps to run, and a progress bar to show how many steps have been executed.

The entry boxes for the time step size and number of steps should be initialized to show the default values for those parameters:  86459 for the time step size and 365 for the number of steps.

After an application has created a RunFrame object named `r` it should be able to call the following methods to interact with `r`:
* `r.dt()` returns the time step size from the time step entry box
* `r.nsteps()` returns the number of time steps indicated in that entry box
* `t.target(n)` sets the upper bound in the progress bar
* `t.update_progress(n)` sets the value to show in the progress bar

Your job is to fill in the body of the `__init__` method so it creates the button, labels, and entry boxes, and fill in the bodies of the methods listed above.

The constructor will be passed a reference to a callback function that will be called when the run button is pushed.  You should pass this function to the constructor that makes the Button.

**Note:**  The run button should be disabled when it is created. The function that loads a data set will enable the button by calling `enable_button` (which has been written for you).

**Note:** &nbsp; In order to test your RunFrame we need to know the names you give to some of the parts of the frame.  Make sure you use the following instance variable names:
* `_run_button` for the button that starts the simulation
* `_dt_entry` for the entry box for time step size
* `_nsteps_entry` for the entry box with the number of steps

##### <span style="color:red">Code:</span>

**Important:** &nbsp; Add your Python code to the following code cell.  Do not delete, rename, or copy this cell.

<!-- grade_id: run_frame, points: 7, grade: true, solution: true -->
```python
class RunFrame(tk.Frame):
    
    def __init__(self, parent, callback):
        tk.Frame.__init__(self, parent)
        
        self['width'] = 200
        self['height'] = 100
        
        self._dt_entry_label = tk.Label(self, text="Time Step Size:")
        self._dt_entry_label.grid(row=0, column=0, sticky=tk.W, padx = 20, pady = 10)

        self._nsteps_label = tk.Label(self, text="Number of steps to run:")
        self._nsteps_label.grid(row=1, column=0, sticky=tk.W, padx = 20, pady = 10)
        
        self._dt_entry = tk.Entry(self)
        self._dt_entry.insert(tk.END,'86459')
        self._dt_entry.grid(row=0, column=1, padx = 10)

        self._nsteps_entry = tk.Entry(self)
        self._nsteps_entry.insert(tk.END,'365')
        self._nsteps_entry.grid(row=1, column=1, padx = 10)
        
        self._run_button = tk.Button(self, text='Run',command=callback,state=tk.DISABLED)
        self._run_button.grid(row=2, column=0, columnspan=2, pady = 5)
        
        self._progress = ttk.Progressbar(self, length=150, maximum=100)
        self._progress.grid(row=3, column=0, columnspan=2, pady = 5)
        
    def dt(self):
        """
        Returns the value entered into the _dt_entry box

        :param self
        :return: the value entered into the _dt_entry box as an integer.
        :rtype: int
        """
        return int(self._dt_entry.get())
    
    def nsteps(self):
        """
        Returns the value entered into the _nsteps_entry box

        :param self
        :return: the value entered into the _nsteps_entry box as an integer.
        :rtype: int
        """
        return int(self._nsteps_entry.get())
        
    def enable_button(self):
        """
        
        """
        self._run_button['state'] = tk.NORMAL
        
    def init_progress(self, n):
        """
        enables the button to use the run button 
        """
        self._progress["maximum"] = n
        
    def update_progress(self, n):
        """
        updates the displayed progress bar by the value of n
        """
        self._progress["value"] = self._progress["value"]+n
        
    def clear_progress(self):
        self._progress['value'] = 0

```

##### <span style="color:red">Tests:</span>

Use the following code cell as a "sandbox" if you want to do your own tests.  You can add additional cells here if you want.

```python

```

##### <span style="color:red">Autograder Tests:</span>

**Important:** &nbsp;  the code cells in this section will be used by `nbgrader` to run automated tests.  Do not move, delete or alter these cells in any way.

```python
def rf_cb():
    global calls
    calls += 1

if IPython:
    rf = RunFrame(tk.Tk(), rf_cb)
    rf.pack()
```

<!-- grade_id: rf_test1, points: 2, grade: true, solution: false -->
```python
if IPython:
    counts = { }
    for x in rf.children.values():
        counts.setdefault(type(x).__name__, 0)
        counts[type(x).__name__] += 1
    assert counts['Button'] == 1
    assert counts['Entry'] == 2
    assert counts['Progressbar'] == 1
    assert rf._dt_entry.get() == '86459'
    assert rf._nsteps_entry.get() == '365'
```

<!-- grade_id: rf_test2, points: 2, grade: true, solution: false -->
```python
if IPython:
    calls = 0
    rf.enable_button()
    rf._run_button.invoke()
    rf._run_button.invoke()
    assert calls == 2
```

<!-- grade_id: rf_test3, points: 2, grade: true, solution: false -->
```python
if IPython:
    rf._dt_entry.delete(0, tk.END)
    rf._dt_entry.insert(0, '100')
    assert rf.dt() == 100
    rf._nsteps_entry.delete(0, tk.END)
    rf._nsteps_entry.insert(0, '1000')
    assert rf.nsteps() == 1000
```

<!-- grade_id: rf_test4, points: 2, grade: true, solution: false -->
```python
if IPython:
    rf.init_progress(100)
    for i in range(5):
        rf.update_progress(10)
    assert rf._progress['value'] == 50
```

##### <span style="color:red">Documentation:</span>

<!-- grade_id: run_doc, points: 5, grade: true, solution: true -->
The process in creating this class involved:

1)adding to the init function for the class by creating the layout for the runframe box that is displayed which includes a run button and a progress bar.

2)create the r.dt() which I created by getting the value entered in the dt entry box and converting it into a usable integer.

3)creating the r.nsteps() method which gets the value entered in the nsteps entry box and converting it into a usable integer.

4)creating the t.target(n) method which sets the upper bound in the progress bar which i did by setting the maximum equal to the given n. 

5)creating the t.update_progress(n) method which sets the value to show in the progress bar and each time it is called adds the value n to its current value.

###  <span style="color:teal">Top Level Application</span>

The application that instantiates all the widgets and links them together is shown below. You do not need to write any code -- the program will work when all of your widget definitions are complete. If you execute this code cell in the Jupyter notebook you will get a working application that will allow you to load a set of planet definitions and run the simulation.

Even though you do not have to write anything you should look closely at this code since it will help you understand the methods you define in each of your components.

To test your GUI, we will export the notebook to a command line application and run the application.  The graders will run the program and check the state of the canvas at the end of the simulation.

#### `load_cb` 

This function is called when the user clicks the Load button.  It uses a dialog to get the name of the file with planet definitions and reads the contents of the file to get a list of TkBody objects.  Notice how:
* the list is passed to the `set_planets` method in the SolarSystemCanvas widget, which will draw circles for all the planets
* the number of bodies is passed to the `reset` method in the ViewControl, which allows the user to click the up/down buttons to change the number of planets displayed

#### `view_cb` 

This function is the callback for the ViewControl component.  Whenever the user changes the value of the counter the new value is passed to `view_planets` so the canvas is updated to show the new number of planets.

#### `run_cb` 

This function is the one that runs the simulation.  It is invoked when the user clicks the Run button in the RunFrame widget.  It gets the simulation parameters from the widget and then calls the `time_step` function to start the simulation.

There body of the `time_step` function could just have a for loop that calls `step_system` for the specified number of time steps.  But since `run_cb` is a callback, Tk won't update the display until the callback is done, which means we won't see the planets move until after the last time step.

The technique implemented here is pretty common in GUI programs.  It runs one time step and then uses the `after_idle` function to schedule a call to run the next time step 0.02 seconds later.  Once the next call is scheduled the callback exits and returns control to the GUI and we will see the planets moving.

```python
root = tk.Tk()
root.title("Solar System")

bodies = None

def load_cb():
    global bodies
    fn = tk.filedialog.askopenfilename()
    bodies = read_bodies(fn, TkPlanet)
    canvas.set_planets(bodies)
    view_counter.set_limit(len(bodies))
    run_frame.enable_button()

def view_cb():
    canvas.view_planets(int(view_counter.get()))
    
def run_cb():
    
    def time_step():
        nonlocal nsteps
        step_system(bodies, dt)
        canvas.move_planets(bodies)
        run_frame.update_progress(1)
        sleep(0.02)
        if nsteps > 0:
            nsteps -= 1
            canvas.after_idle(time_step)
        else:
            run_frame.clear_progress()
        
    nsteps = run_frame.nsteps()
    run_frame.init_progress(nsteps)
    dt = run_frame.dt()
    canvas.after_idle(time_step)

canvas = SolarSystemCanvas(root)
canvas.grid(row = 0, column = 0, columnspan = 3, padx = 10, pady = 10, sticky="nsew")

tk.Button(root, text='Load', command=load_cb).grid(row=1, column=0, pady = 20)

view_frame = tk.Frame(root, width=100)
tk.Label(view_frame, text='Planets to View').pack()
view_counter = Viewbox(view_frame, view_cb)
view_counter.pack()
view_frame.grid(row=1, column=1, pady=20)

run_frame = RunFrame(root, run_cb)
run_frame.grid(row=1, column=2, pady=20)

if Main and not IPython:
    try:
        bodies = read_bodies("solarsystem.txt", TkPlanet)
        canvas.set_planets(bodies)
        view_count.reset(len(bodies))
        for i in range(5):
            view_count._spinbox.invoke('buttondown')
        run_frame._nsteps_entry.delete(0, tk.END)
        run_frame._nsteps_entry.insert(0,'100')
        root.update()
        run_frame._run_button.invoke()
    except Exception as err:
        print(err)
    input('hit return to continue...')

```

##### <span style="color:red">Simulation Score</span>

**Important:** &nbsp; You do not have to write anything for this part of the project.  Leave the cell below empty, in spite of the fact that is says "your documentation here".

The graders will use the cell to write comments about your GUI and to enter the grade for your solar system simulation.  **Do not edit, move, rename, or delete the following cell.**

<!-- grade_id: simulation, points: 20, grade: true, solution: true -->
YOUR DOCUMENTATION HERE